# Problem Set for Python Videos

Both video series should be viewed before starting this problem set.

This problem set is meant to be rather challenging. This problem set is meant to stretch your skills so you are ready for the first day of class. Do not get discouraged if you can not complete all of the problem set, just do what you can and that will be sufficient. The assert cells after each problem can be used to test your solution.

Although this problem set is ungraded, please submit it on courseworks as HW0 for our review. 

Finally, please complete the following short survey after you have completed this problem set: https://goo.gl/forms/Fi4QxauZU4RXVVX62

---

The integers above zero and below 10 that are multiples of 3 or 7 are 3, 6, 7, 9. The product of these numbers is 1134.

Write a function named `sparse_factorial` to find the product of all integers above zero and below `x` that are multiples of 3 or 7.

<!-- grade_id: q1-solution, points: 0, grade: false, solution: true -->
```python
def sparse_factorial(x):
    # YOUR CODE HERE
    product = 1
    for num in range(1,x+1):
        if (num%3==0 or num%7==0):
            product=product*num
    return product
    print(product)
    # raise NotImplementedError()
```

<!-- grade_id: q1-tests, points: 1, grade: true, solution: false -->
```python
assert sparse_factorial(10) == 1134
assert sparse_factorial(50) == 4238458163289827194798080000
```

---

A palindromic number reads the same both ways. For example, 909 and 4321234. Write a function to determine if a number is palindromic. The function should be called `is_palindromic` and should return `True` or `False`.

<!-- grade_id: q2-solution, points: 0, grade: false, solution: true -->
```python
def is_palindromic(x):
    # YOUR CODE HERE
    if str(x) == str(x)[::-1]:
        return True
    else:
        return False   
    
   # raise NotImplementedError()
```

<!-- grade_id: q2-tests, points: 1, grade: true, solution: false -->
```python
assert not is_palindromic(10003)
assert is_palindromic(4321234)
```

---

Write a function named `is_sorted`. This function should take an iterable and should return True if all items in the iterable are sorted in order from smallest to largest. Hint: you can represent infinity with `float("inf")` in Python.

<!-- grade_id: cell-52b7096ac1df7300, points: 0, grade: false, solution: true -->
```python
def is_sorted(iterable):
    # YOUR CODE HERE
    m = iterable.sort()
    if iterable == m:
        return True
    else:
        return False
    #raise NotImplementedError()
```

<!-- grade_id: cell-153652db4189a98e, points: 1, grade: true, solution: false -->
```python
assert is_sorted([1, 2, 3])
assert not is_sorted([3, 2, 1])
```

---

Write a function named `phone_format` that receives a dictionary of phone number information and returns the formatted phone number as a string.

```
{
    'area_code': 510,
    'prefix': 347,
    'line_number': 7446,
}

===> '(510) 347-7446'
```

<!-- grade_id: cell-2da8f8025e4366b4, points: 0, grade: false, solution: true -->
```python
def phone_format(dict_):
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-99ace83e0a9237b2, points: 1, grade: true, solution: false -->
```python
assert phone_format({'area_code': 510, 'prefix': 347, 'line_number': 7446}) == '(510) 347-7446'
assert phone_format({'area_code': 947, 'prefix': 746, 'line_number': 8664}) == '(947) 746-8664'
```

---

Create a function named `cast`. This function should take a value and an optional type. If a type is not given, the value should be returned unchanged. If a type is given, the value should be cast to that type and returned.

<!-- grade_id: cell-2c7d16faf74a2acd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cast(val,type_=None):
    if type_ == None:
        return val
    elif type_ == str:
        return str(val)
    else:
        return complex(val)

#raise NotImplementedError()
```

<!-- grade_id: cell-db52dc3294e9f6aa, points: 1, grade: true, solution: false -->
```python
assert cast(1) is 1
assert cast(1, type_=complex) == (1+0j)
assert cast(1, type_=str) == '1'
```

---

The following number has a series of 5 consecutive digits for which the sum of those digits is greater than the sum of any other series of 5 consecutive digits in the number. 

    314159265358979323846264338327950288419716939937510

    3141592653|58979|323846264338327950288419716939937510
    
    5 + 8 + 9 + 7 + 9 = 38

Write a function `slice_search` to look for an abitrarily large series of digits with a maximum sum and find the 0-based index of the first of those digits. A 0-based index is one where the first item in the sequence has an index of 0. Ie. 3 would have the index of 0 in this number, 1 would have the index of 1, and 4 would have the index of 2:

    314
    
If there are multiple series that share the same maximum sum, return the lowest index.

<!-- grade_id: cell-2a90e68b8a21714b, points: 0, grade: false, solution: true -->
```python
def slice_search(number, digit_count):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e756e300a04b91e1, points: 1, grade: true, solution: false -->
```python
assert slice_search(314, 1) == 2
assert slice_search(314159265358979323846264338327950288419716939937510, 5) == 10
assert slice_search(3728383, 2) == 3
```

---

Write a function named `set_overlap_maximizer` that takes a list of lists. The function should find the overlap of all possible parings of lists and find the two lists that have the greatest overlap of values. The function should return a 2-tuple (a tuple with two items in it) with the indexes of the two lists in the outer list. The 2-tuple should be ordered in ascending order (ie. (1, 2), not (2, 1). If there are multiple pairs that share the same maximum overlap, return the tuple with the lowest indexes.

<!-- grade_id: cell-ffed39796700ce8b, points: 0, grade: false, solution: true -->
```python
def set_overlap_maximizer(list_of_lists):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-247b92d15cfda6bb, points: 1, grade: true, solution: false -->
```python
assert set_overlap_maximizer([[1, 2, 3], [3, 4, 5], [4, 5, 6]]) == (1, 2)
assert set_overlap_maximizer([[True, False, None, 8, 4, 'a'], ['a', None, False], [8, 9, 2]]) == (0, 1)
```

---

A leap year occurs on any year evenly divisible by 4, but not on a century unless it is divisible by 400.
Write a function named `is_leap_year` which takes a year greater than 0 and returns True if a that year is a leap year. For fun, find the percentage of years between 0 and 2018 (inclusive of 2018) that are leap years. 

<!-- grade_id: cell-28812396fce16207, points: 0, grade: false, solution: true -->
```python
def is_leap_year(year):
    # YOUR CODE HERE
    if year%400==0 or year%4==0:
        return True
    else:
        return False
    
def print_leap_percent(year):
    count_leaps = 0
    for num in range(0,year+1):
        if is_leap_year(num)==True:
            count_leaps+=1
        else:
            count_leaps=count_leaps
    print("For years between 0 and "+str(year)+", "+str(round((100*count_leaps/year),2))+"% are leap years")
    #raise NotImplementedError()
```

```python
print_leap_percent(2018)
```

```python
assert not is_leap_year(1900)
```

<!-- grade_id: cell-28876b241a449a1b, points: 1, grade: true, solution: false -->
```python
assert is_leap_year(2000)
assert is_leap_year(2004)
assert not is_leap_year(1900)
assert not is_leap_year(2003)
```

---

You are delievered a file that has several hundered lines of text. Lines are separated by a single new line character while paragraphs are separated by two new line characters. For example

    This is one line.
    This is another line in the same paragraph.
    
    This is a new line in a new paragraph.
    
Write a function named `paragrapher` that takes a single blob of text and returns a tuple of paragraphs of text. The paragraphs should not contain new line characters and the text should be in the same order as it is received. 

<!-- grade_id: cell-a03e58f0329956ec, points: 0, grade: false, solution: true -->
```python
def paragrapher(text):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6e1b32344496a47f, points: 1, grade: true, solution: false -->
```python
text = """
This is one line.
This is another line in the same paragraph.

This is a new line in a new paragraph.
"""
expected_result = (
    'This is one line. This is another line in the same paragraph.',
    'This is a new line in a new paragraph.'
)
assert paragrapher(text) == expected_result
```

---

Write a function named `identity` that produces identity matrices of size N by N. The matrices should be made of tuples within tuples. The function should have a docstring that states "Return an identity matrix of n by n size."

<!-- grade_id: cell-b1d2aa9d5c36ec74, points: 0, grade: false, solution: true -->
```python
def identity(n):
    # YOUR CODE HERE
    """Return an identity matrix of n by n size."""

    raise NotImplementedError()
```

<!-- grade_id: cell-1602d23627688ffc, points: 1, grade: true, solution: false -->
```python
assert identity(2) == ((1, 0), (0, 1))
assert identity(3) == ((1, 0, 0), (0, 1, 0), (0, 0, 1))
assert identity.__doc__.strip() == 'Return an identity matrix of n by n size.'
```

---

The `random` library has several functions generating random values. Write a function named `count_attempts` that takes a number `x` between 0 and 1 and returns the number of times `random.random` had to be called before it produced a number within 0.01 of `x`.

<!-- grade_id: cell-6178408758bd9ac6, points: 0, grade: false, solution: true -->
```python
import random

def count_attempts(x):
    # YOUR CODE HERE
    count_times = 0
    num = random.random()
    
    while abs(x-num) > 0.01:
        count_times +=1
    print(count_times+1)
    
    #raise NotImplementedError()
```

<!-- grade_id: cell-92634588c49be1bf, points: 1, grade: true, solution: false -->
```python
random.seed(42)

assert count_attempts(0.6) == 30
assert count_attempts(0.93) == 24
```

---

Create a Ceaser cipher. Implement a function `ceaser_encrypt` that takes some text (`plaintext`) and returns ciphertext. To compute the ciphertext, iterate through each letter in the plain text and replace it with the letter that is 13 characters forward in the ASCII alphabet. For letters that are too far towards the end of the alphabet, start back at the beginning of the alphabet and continue counting up to 13. The function should convert `plaintext` to lower case before operating on the text. All punctuation should be ignored by the cipher.

For example:

```python
>>> def ceaser_encrypt(plaintext):
...     # implement me
...     pass
... 
>>> ceaser_encrypt('You rock!')
'lbh ebpx!'
```

Finally, what is the result if you pass text encrypted by `ceaser_encrypt` back into `ceaser_encrypt`?

<!-- grade_id: cell-ef3a757c826739d5, points: 0, grade: false, solution: true -->
```python
from string import ascii_letters

def ceaser_encrypt(plaintext):
    # YOUR CODE HERE
    plaintext=plaintext.lower()
    output_string = ''
    for character in plaintext:
        if character in ascii_letters:
            output_string = output_string + ascii_letters[ascii_letters.index(character)+13]
        else:
            output_string +=character
    print(output_string.lower())
            
        
    #raise NotImplementedError()
```

<!-- grade_id: cell-5bf52ef2852c8eb2, points: 1, grade: true, solution: false -->
```python
text = 'Secrets secrets!'
assert ceaser_encrypt(text) == 'frpergf frpergf!'
assert ceaser_encrypt(ceaser_encrypt(text)) == text.lower()
```

---

Write a function named `factorial` to find the factorial for a given number.

```
>>> factorial(3)
6
```

<!-- grade_id: cell-3296fbc3c20c379d, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    # YOUR CODE HERE
    if n<0:
        print("No factorial for negatives")
    elif n==0:
        print(1)
    else:
        factorial = 1
        for num in range(1,n+1):
            factorial = factorial*num
        return(factorial)
    #raise NotImplementedError()
factorial(50)
```

Output:
```
30414093201713378043612608166064768844377641568960512000000000000
```

<!-- grade_id: cell-67af2acae157a971, points: 1, grade: true, solution: false -->
```python
assert factorial(3) == 6
assert factorial(50) == 30414093201713378043612608166064768844377641568960512000000000000
```

---

Create a function named `change` that returns a `dict` containing the number of dollars, quarters, dimes, nickels, and pennies required for a given amount of change. Bigger denominations should be use before smaller ones are used. You can assume you are in a utopia were you can treat rounding errors in pennies by rounding up to the nearest penny. 

An example of the function:

```
>>> get_change_count(0.92)
{
    'quarter': 3,
    'dime': 1,
    'nickel': 1,
    'penny': 2,
}
```

<!-- grade_id: cell-bdfb5e93ff922394, points: 0, grade: false, solution: true -->
```python
def change(amount):
    # YOUR CODE HERE - as far as i got but i know not correct
    dollar, quarter, dime, nickel, penny = 1, .75, .1, .05, .01
    coins = (dollar, quarter, dime, nickel, penny)
    change_count = []
    for i in coins:
        while amount >= i:
            change_count.append(i)
            amount=amount-i
    print(change_count)
print(change(1.92))   
    #raise NotImplementedError()
```

Output:
```
[1, 0.75, 0.1, 0.05, 0.01]
None
```

<!-- grade_id: cell-5f73a8f88bc021f9, points: 1, grade: true, solution: false -->
```python
assert change(1.92) == {'dime': 1, 'dollar': 1, 'nickel': 1, 'penny': 2, 'quarter': 3}
assert change(0.13) == {'dime': 1, 'dollar': 0, 'nickel': 0, 'penny': 3, 'quarter': 0}
```

---

Create a function named `count` that takes a string of characters, counts each occurrence of each character, and produces a compressed string with each character and its count. For example:

`'aabbccababcc'` would compress to `'a4b4c4'`

The result string should be sorted.

<!-- grade_id: cell-3137e1898c9f948c, points: 0, grade: false, solution: true -->
```python
import collections

def count(string):
    output = ''
    for char in string:
        if char in output:
            

        
    raise NotImplementedError()
```

<!-- grade_id: cell-b127e8ae57db64cc, points: 1, grade: true, solution: false -->
```python
assert count('aabbccababcc') == 'a4b4c4'
assert count(reversed('aabbccababcc')) == 'a4b4c4'
```

```python

```