**Q**: Create a object that represents a store. Have the initialization method take the square footage of the store. Create a method that can calculate the electrical costs for the store given:

- $0.15 per kWh
- kWh per square foot per hour of 0.19

Ensure your class has the following API:

Instance attributes: 
- `square_footage`

Class level attributes:
- DOLLARS_PER_KWH
- KWH_PER_SQUARE_FOOTAGE_PER_HOUR

[2 points]

<!-- grade_id: cell-e0ff639912440f53, points: 0, grade: false, solution: true -->
```python
class Store:
    DOLLARS_PER_KWH = 0.15
    KWH_PER_SQUARE_FOOTAGE_PER_HOUR = 0.19
    
    def __init__(self, sf):
        self.square_footage = sf
        
    def electrical_cost(self,ec):
        self.electrical_cost = ec
        
        return int(
            round(
                (Store.DOLLARS_PER_KWH*Store.KWH_PER_SQUARE_FOOTAGE_PER_HOUR*self.square_footage*self.electrical_cost),
                1))
        
```

<!-- grade_id: cell-e4863bba06d08e3f, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store(2000).square_footage == 2000
assert not hasattr(Store, 'square_footage')
### END TESTS
```

<!-- grade_id: cell-47704c4e3e538b07, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert hasattr(Store, 'DOLLARS_PER_KWH')
assert hasattr(Store, 'KWH_PER_SQUARE_FOOTAGE_PER_HOUR')
### END TESTS
```

<!-- grade_id: cell-e9ab2e3e055a711f, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert Store(2000).electrical_cost(900) == 51300
### END TESTS
```

---

**Q**: Create a counter class that takes a string and can produce (via a method named count) a `dict` containing each word in the string as its keys and the counts of each word in the string as its values. Words are any token separated by a space. The class should not be case sensitive (ie. "It" and "it" are the same word).


```python
>>> c = Counter('This is my best string. It is a great string.')
>>> c.count()
{'a': 1,
 'best': 1,
 'great': 1,
 'is': 2,
 'it': 1,
 'my': 1,
 'string.': 2,
 'this': 1}
 ```
 
**Class API**

Instance attributes:
  - "string"
  
You can not use the collections module.
 
[2 points]

<!-- grade_id: cell-0be5d346f805f3fb, points: 0, grade: false, solution: true -->
```python
class Counter:
    def __init__(self,string):
        self.string = string.lower().split()
    def count(self):
        frequency_dict={}
        for word in self.string:
            frequency_dict[word]=0
            frequency_dict[word]=self.string.count(word)           
        return frequency_dict
```

<!-- grade_id: cell-59639940597f3439, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert len(counter.count()) == 8
### END TESTS
```

<!-- grade_id: cell-ea0b8e76389a0ae2, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert hasattr(counter, 'string')
assert not hasattr(Counter, 'string')
### END TESTS
```

<!-- grade_id: cell-91fa2691c9bea2eb, points: 1, grade: true, solution: false -->
```python

```

---

**Q**: Update the counter class to have the appropriate methods for presentation when being passed to `repr` and `str`. `repr` should return the object class name and a truncated version of the counter's string. The string should be truncated to 10 characters and an ellipsis added. `str` should return the object class name and the key count. [1 point]

```python
>>> obj = Counter('Fe Fi Fo Fum, where is my drum')
>>> repr(obj)
<Counter 'Fe Fi Fo Fum, ...'>
>>> str(obj)
'Counter(8)'
```

<!-- grade_id: cell-eff92fe808b528d9, points: 0, grade: false, solution: true -->
```python
class Counter:
    def __init__(self,string):
        self.string = string.lower().split()
        self.string_og = string
    def count(self):
        frequency_dict = {}
        for word in self.string:
            frequency_dict[word]=0
            frequency_dict[word]=self.string.count(word)           
        return frequency_dict
    def ___repr__(self):
        return "<Couner '%s...'>" % (self.string_og[:10])
    def __str__(self):
        frequency_dict = {}
        for word in self.string:
            frequency_dict[word]=0
            frequency_dict[word]=self.string.count(word)  
        return 'Counter(%s)' % (str(len(frequency_dict)))
a = Counter('This is my best string. It is a great string.')
print(repr(a))
```

<!-- grade_id: cell-42a5d9c016cae56f, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-76d7e759c0c50b24, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert repr(counter) == "<Counter 'This is my...'>"
### END TESTS
```

<!-- grade_id: cell-d5d109b50b154b18, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
counter = Counter('This is my best string. It is a great string.')
assert str(counter) == 'Counter(8)'
### END TESTS
```

**Q**: Create a custom exception for a severe lack of parties. It should be called `SevereLackOfPartiesError` and should subclass `ValueError`. The class should have a docstring that describes what the error represents. Use your imagination for the exact text. 

[1 point]

<!-- grade_id: cell-7368367ae633aef9, points: 0, grade: false, solution: true -->
```python
class SevereLackOfPartiesError(ValueError):
    """Dangerous levels of lame detected"""
    pass
```

<!-- grade_id: cell-78f404be90d25d2e, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-1b63d2e6166363a1, points: 0.5, grade: true, solution: false -->
```python

```

---

**Q**: Create a class `Str` that subclasses `str`. Add a method to the subclass that checks if the string does not start with a given string. Call that method `notstartswith`. This class should only take up 3 lines of code however you will not be graded on the length of your class.

[1 point]

<!-- grade_id: cell-931411951b8973e6, points: 0, grade: false, solution: true -->
```python
class Str(str):
    def notstartswith(self,x):
        return self.startswith(x)
```

<!-- grade_id: cell-dde9e6cf99caf6bc, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-767d16e4e70c6601, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7f7773df3332732e, points: 0.25, grade: true, solution: false -->
```python

```

----

**Q**: Create a `BankAccount` class that is instantiated with an account number and an optional initial balance. Add methods to withdraw and deposit to the account's balance. Also, add a method that allows for the casting of objects of this class to a string.

    a = BankAccount(12345678, balance=10)
    a.deposit(5)
    a.withdraw(3)
    str(a)

    'Account ID: 12345678 - Balance: 12'

[2 points]

<!-- grade_id: cell-31e546afd67a9412, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    def __init__(self,account_id,balance=0):
        self.account_id = account_id 
        self.balance = balance
        
    def withdraw(self, amount):
        self.balance = self.balance - amount
    
    def deposit(self,amount):
        self.balance = self.balance + amount
    
    def __str__(self):
        return 'Account ID: '+ str(self.account_id) + ' - Balance: ' + str(self.balance) 
```

<!-- grade_id: cell-78f78a2ea5d08b99, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert str(BankAccount(4839283, balance=2010)) == 'Account ID: 4839283 - Balance: 2010'
### END TESTS
```

<!-- grade_id: cell-e706bb44eae7db75, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert BankAccount(1).balance == 0
### END TESTS
```

<!-- grade_id: cell-da11f9aeadad2c39, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = BankAccount(1)
a.deposit(150)
assert a.balance == 150
### END TESTS
```

<!-- grade_id: cell-0f6bc2f6bd8c7a9c, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = BankAccount(1)
a.deposit(150)
a.withdraw(27)
assert a.balance == 123
### END TESTS
```

----

**Q**: Subclass the `BankAccount` class from the previous question to create a `MinBalanceBankAccount` class. This class should require a minimum balance at all times and should raise an exception (of the type `ValueError`) when an attempt to drop the balance below the minumum occurs. The `MinBalanceBankAccount` should be instantiated with the account number, inital balance, and the minimum balance. No error checking is necessary on instantiation.

```
a = MinBalanceBankAccount(account_id=1, balance=100, minimum_balance=10)
a.withdraw(100)
...
ValueError: Sorry, minimum balance must be maintained.
```


[2 points]

<!-- grade_id: cell-ffede70c71976fef, points: 0, grade: false, solution: true -->
```python
class MinBalanceBankAccount(BankAccount):
    def __init__(self, account_id, balance, minimum_balance):
        self.account_id = account_id
        self.balance = balance
        self.minimum_balance = minimum_balance
    
    def withdraw(self,amount):
        if self.balance - amount < self.minimum_balance:
            raise ValueError('Sorry, minimum balance must be maintained.')
        else:
            self.balance = self.balance - amount
```

<!-- grade_id: cell-4d189116d6cb3d51, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
try:
    MinBalanceBankAccount(1, 100, 10).withdraw(100)
    assert False, 'Should have raised value Error'
except ValueError as e:
    assert str(e) == 'Sorry, minimum balance must be maintained.'
### END TESTS
```

<!-- grade_id: cell-496de9ecb2b38244, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-74b6dba911c7df84, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert issubclass(MinBalanceBankAccount, BankAccount)
### END TESTS
```

<!-- grade_id: cell-3d7deeb06f900ce3, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a = MinBalanceBankAccount(1, 100, 10)
a.deposit(150)
a.withdraw(27)
assert a.balance == 223
### END TESTS
```

---

**Extra Practice** Create a class that produces objects that are both iterables and iterators. 

<!-- grade_id: cell-acfb6cbf62089d2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```