<!-- grade_id: cell-e5713946811b5d52, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import subprocess
import tempfile
import os


def clean_bash_function(text):
    if 'rm' in text:
        raise ValueError('rm found in text of script. No points.')
    
    text = text.strip()
    # Get first \n
    lines = text.split('\n')

    return '\n'.join(line.lstrip() for line in lines)


def run_script(text, args=None, timeout=None):
    args = args or []
    text = clean_bash_function(text)
    with tempfile.NamedTemporaryFile() as file:
        file.write(text.encode())
        file.flush()
        os.chmod(file.name, 0o777)    
        response = subprocess.run(
            ['/bin/bash', file.name] + args, 
            timeout=timeout, 
            stdout=subprocess.PIPE, 
            stderr=subprocess.PIPE,
        )
        return response
### END TESTS
```

**DO NOT INCLUDE THE STRING  `rm`  ANYWHERE IN YOUR SCRIPTS OR YOU COULD GET ZERO POINTS FOR THE QUESTION.**

**Q**: Write a function `get_bash_function` that returns the text of a bash script. For example:

```python
def get_bash_function():
    return """
    #!/bin/bash
    echo Hello!
    """
```

The bash script should not return "hello" but should count down from 5 to 1 (inclusive), printing each number as it goes and should print "GO!" after printing 1. All values should be printed to their own line. All values should be printed to standard out. The script should pause for one second in between printing each line. A `for` loop should be used in the count down. If you script were run at the command line, output would be this:

```bash
$ /bin/bash your-script.sh
5
4
3
2
1
GO!
```

[2 points]

<!-- grade_id: cell-40a07e3e81c26639, points: 0, grade: false, solution: true -->
```python
def get_bash_function():
    return """
    #!/bin/bash
    i=5
    while (( i >= 1 ))
    do
        echo $(( i-- ))
        sleep 1s
    done
    echo "GO!"
    """
```

<!-- grade_id: cell-ae95381e40692307, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
text = get_bash_function()
response = run_script(text, timeout=10)
assert response.stdout.decode().strip().split() == ['5', '4', '3', '2', '1', 'GO!']
### END TESTS
```

<!-- grade_id: cell-de51c6a3de2b41e7, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-396d56170dc6809e, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-e562b568261ce8ae, points: 0.5, grade: true, solution: false -->
```python

```

---

**Q**: Write a function `get_bash_function` that returns the text of a bash script. For example:

```python
def get_bash_function():
    return """
    #!/bin/bash
    echo Hello!
    """
```

The bash script should not return "hello" but should take one argument and use it within a greeting. Please use the greeting seen in the example. All output should be printed to standard out. If you script were run at the command line, output would be this:

```bash
$ /bin/bash your-script.sh Paul
Well, hello Paul! Nice to meet you today :)
```
[1 point]

<!-- grade_id: cell-0ae8eeb9e364516c, points: 0, grade: false, solution: true -->
```python
def get_bash_function():
    return """
    #!/bin/bash
    echo Well, hello $@! Nice to meet you today :)
    """
```

<!-- grade_id: cell-2af32a0dc3f329d1, points: 1, grade: true, solution: false -->
```python

```

---

**Q**: Write a function `get_bash_function` that returns the text of a bash script. For example:

```python
def get_bash_function():
    return """
    #!/bin/bash
    echo Hello!
    """
```

The bash script should not return "hello" but should take one argument and find that argument's factorial. All output should be printed to standard out. If you script were run at the command line, output would be this:

```bash
$ /bin/bash your-script.sh 5
120
```
[2 points]

<!-- grade_id: cell-fd417f3e5c08b2a1, points: 0, grade: false, solution: true -->
```python
def get_bash_function():
    return """
    #!/bin/bash
    output=1
    for num in $(seq $@)
    do
        output=$((output*num))
    done
    echo $output
    """
```

<!-- grade_id: cell-3b22872326e5c128, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
text = get_bash_function()

response = run_script(text, args=['5'], timeout=2)

assert response.stdout.decode().strip() == '120'
### END TESTS
```

<!-- grade_id: cell-7528b05a98e71db9, points: 1, grade: true, solution: false -->
```python

```