**Q**: Explain why there are only two items in the evaluated set `{1, True, 1.0, False, 0, None}`. [1 point]

<!-- grade_id: cell-a6e7085bd404dd68, points: 1, grade: true, solution: true -->
The only three items that are evaluated are 1, True, and 1.0 beause they all have values of non zero. False, 0, and None are all zeros.

---

**Q**: Create a function `evens` that returns all the even numbers between 0 and a given number x. Zero should be included in the result but x should not be included unless x is zero. [1 point]

<!-- grade_id: cell-f178c23df522f905, points: 0, grade: false, solution: true -->
```python
def evens(x):
    """Returns list of even numbers between 0 and given number"""
    output=[num for num in range(0,x) if num%2==0]
    return output

    raise NotImplementedError()
```

<!-- grade_id: cell-80b80f97d1d635e1, points: 1, grade: true, solution: false -->
```python

```

---

**Q**: Implement Fizz-Buzz-Bang. Write a program that returns the number given. But for multiples of two return "Fizz" instead of the number, for the multiples of three return "Buzz", and for the multiples of four return "Bang". For numbers which are multiples of both two and three return "FizzBuzz". For numbers which are multiples of both two, three, and four return "FizzBuzzBang". [2 points]

<!-- grade_id: cell-fda6d5a34cfe1fec, points: 0, grade: false, solution: true -->
```python
def fizz_buzz_bang(i):
    """Returns FizzBuzzBang combo based on number entered"""
    if i%2==0 and i%3==0 and i%4==0:
        return "FizzBuzzBang"
    elif i%2==0 and i%3==0:
        return "FizzBuzz"
    elif i%2==0:
        return "Fizz"
    elif i%3==0:
        return "Buzz"
    elif i%4==0:
        return "Bang"
    else:
        return i
        
    raise NotImplementedError()
```

<!-- grade_id: cell-eb510034e9d6fb63, points: 2, grade: true, solution: false -->
```python

```

---

**Q**: Write a function `compress` that takes a string of characters, counts each occurrence of each character, and produces a compressed string with each character and its count. For example:

`'aaaabbbbcccc'` would compress to `'a4b4c4'`

You can assume that the source string will be sorted. The result string should be sorted. [2 points]

<!-- grade_id: cell-63c10b65c17caab0, points: 0, grade: false, solution: true -->
```python
def compress(letter_string):
    """Returns compressed string with letter frequency counts"""
    output=''
    for char in letter_string:
        if char not in output:
            output=output+str(char)+str(letter_string.count(char))
    return output

    raise NotImplementedError()
```

<!-- grade_id: cell-33ec73d3be0d6640, points: 2, grade: true, solution: false -->
```python

```

----

**Q**: Write a function `endpoints` that takes a list of numbers (eg. `[5, 10, 15, 20, 25]`) and returns a new list of only the first and last elements of the given list (eg. `[5, 25]`). The function should return an empty list if an empty list is passed in. The function should **not** use any variables besides the passed in argument `list_`.  [1 point]

<!-- grade_id: cell-e06b186eadc3d5c3, points: 0, grade: false, solution: true -->
```python
def endpoints(list_):
    """Returns 1st and last elements of list"""
    if len(list_)==0:
        return list_
    
    else:
        output = list_[0], list_[-1]
        return list(output)
    
    raise NotImplementedError()
```

<!-- grade_id: cell-f7e7bb26b30ed92a, points: 1, grade: true, solution: false -->
```python

```

---

**Q**: Write a function `common` that consumes two lists and that produces a list that contains only the common elements (without duplicates). Make sure your program works on two lists of different sizes and returns the common elements in sorted ascending order.


```python
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]

# result
[1, 2, 3, 5, 8, 13]
```

[1 point]

<!-- grade_id: cell-1ab667ab32f392f9, points: 0, grade: false, solution: true -->
```python
def common(a, b):
    """Return list of common elements between two lists"""
    interim = [item for item in a if item in b]
    output=[]
    
    for num in interim:
        if num not in output:
            output.append(num)
    return output

    raise NotImplementedError()
```

<!-- grade_id: cell-8cd8aa2c0c722e82, points: 1, grade: true, solution: false -->
```python

```

----

**Q**: Create a function `search` that searches a dictionary's keys for a keyword and returns a set of the associated values for all keys that contain the keyword. If it helps, assume the dictionary is keyed by a few of your friends' first names and valued by their associated favorite color. Make sure the search is case-insensitive.

```python
d = {
    'Urbi': 'Green',
    'Derek': 'Blue',
    'Jon': 'Yellow',
    'Lin': 'Purple',
}
```
- Searching for 'j' results in {'Yellow'}
- Searching for 'i' results in {'Green', 'Purple'}

[2 points]


<!-- grade_id: cell-71b5d64fa7de099d, points: 0, grade: false, solution: true -->
```python
def search(dict_, search_term):
    """Returns values where search term appears in key"""
    search_term = search_term.lower()
    dict_={key.lower():value for key,value in dict_.items()}
    output = {value for key,value in dict_.items() if search_term in key}
    
    return output

    raise NotImplementedError()

```

<!-- grade_id: cell-f6efe948113bdc7a, points: 2, grade: true, solution: false -->
```python

```

----

**Q**: Write a script that appends 100,000 items to a an empty list; the item can be anything so the number 1 would do. Time how long it takes (with a stop watch on the Internet or your phone). Write a script that prepends 100,000 items to an empty list and time how long it takes. 

What script takes longer and why? No more than 50 words. 

If either script runs for more than a minute, stop it and assume that it ran infinitely long. [1 point]

<!-- grade_id: cell-670e7107d58db7bf, points: 0, grade: true, solution: true -->
```python
def big_list_b():
    """Appends 'z' to list"""
    list_a = []
    for count in range(0,100001):
        list_a.append('z')
    return list_a

    raise NotImplementedError()
```

<!-- grade_id: cell-c1f537f3db59f76e, points: 0, grade: true, solution: true -->
```python
def big_list_b():
    """Prepends 'z' to a list"""
    list_b = []
    for count in range(0,100001):
        list_b.insert(0,'z')
    
    return list_b

    raise NotImplementedError()
big_list_b()
```

<!-- grade_id: cell-91f1edcf7c372b7c, points: 1, grade: true, solution: true -->
Perpending took 4s whereas appending was nearly instant. This is because with prepending all the positions of the characters currently held in the list need to be 'moved down' a spot to make room for the new character at position 0. With appending, the new character is simply placed at the end of the len(list_b)

---

**Q**: Create a function `get_lambda` that returns a `lambda` function that returns a tuple containing N instances of a given value.

```python
>>> f = get_lambda()
>>> f(4, True)
(True, True, True, True)
```

[1 point]

<!-- grade_id: cell-440167c0cf8ec924, points: 0, grade: false, solution: true -->
```python
def get_lambda(instance, given):
    """Returns a tuple containing N instances of a given value"""
    output = map(lambda given: given*instance,given)
    
    return tuple(output)
                   
    raise NotImplementedError()

get_lambda(4,'v')
```

Output:
```
('vvvv',)
```

<!-- grade_id: cell-d1f15d8c975bef1e, points: 1, grade: true, solution: false -->
```python

```

**Q**: Create a function `get_lambda` that returns a `lambda` function that takes a tuple and raises each element in the tuple to the second power.

```python
>>> f = get_lambda()
>>> f((1, 2, 3, 4))
(1, 4, 9, 16)
```
[1 point]

<!-- grade_id: cell-70896f7d7cc69cde, points: 0, grade: false, solution: true -->
```python
def get_lambda(givens):
    """Raises each element in tuple to second power"""
    return tuple(map(lambda x: x**2, givens))
    
    raise NotImplementedError()
```

<!-- grade_id: cell-7b39c52cfbb8fe3a, points: 1, grade: true, solution: false -->
```python

```

---

**Q**: Write a function `get_cubes` that builds a new list via a list comprehension. The list comprehension should contain in sorted order, all cubes for numbers between 0 a given number `x` if the number is divisible by 3.

```python
>>> get_cubes(10)
[0, 27, 216, 729]
```

[1 point]

<!-- grade_id: cell-70c782eb764fa9f9, points: 0, grade: false, solution: true -->
```python
def get_cubes(x):
    """Returns list of cubes for numbers in range that are divisible by 3"""
    output=[num**3 for num in range(0,x+1) if num%3==0]
    
    return output

get_cubes(10)
```

Output:
```
[0, 27, 216, 729]
```

<!-- grade_id: cell-f25ee1a76203c90e, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7f2062eecdd55c4f, points: 1, grade: true, solution: false -->
```python

```

---

**Q**: The function calls `random.choice(string.ascii_lowercase)` return a random lowercased ASCII letter. Create a function `get_count` that returns the number of unique letters returned by `x` sequential calls to these functions. Your function should use a set comprehension.

```python
>>> get_count(10)
8
>>> get_count(10)
7
>>> get_count(10)
9
>>> get_count(7)
5
```

[1 point]

<!-- grade_id: cell-00e884959ab1fddd, points: 0, grade: false, solution: true -->
```python
import string
import random

def get_count(x):
    """Returns the number of unique letters after calling function x times"""
    letters_generated = ''
    unique = []
    
    for i in range(x+1):
        letter = random.choice(string.ascii_lowercase)
        letters_generated+=letter
      
    for char in letters_generated:
        if char not in unique:
            unique.append(char)
            
    return len(unique)
    
    raise NotImplementedError()
    
get_count(10)
```

<!-- grade_id: cell-06fffca9a7420049, points: 1, grade: true, solution: false -->
```python

```

---

**Q**: Create a function `get_debits` that takes a dictionary and returns a new dictionary were each key in the given dictionary is prefixed with "debit\_" and each amount's value is multiplied by -1. The function should make use of a dictionary comprehension.


```python
>>> d = {'helicopter': 430, 'hanger': 75}
>>> get_debits(d)
{'debit_hanger': -75, 'debit_helicopter': -430}
```

[1 point]

<!-- grade_id: cell-225b03fcb668cf76, points: 0, grade: false, solution: true -->
```python
def get_debits(d):
    """Returns dictionary of debits"""
    output_dict={'debit_'+str(key):value*-1 for (key,value) in d.items()}
    
    return output_dict
    
    
    raise NotImplementedError()

d={'helicopter':430, 'hanger':75}
get_debits(d)
```

<!-- grade_id: cell-4ad3da6ea4b02b76, points: 1, grade: true, solution: false -->
```python

```

----

**BONUS Q**: Create a function that returns a list that contains itself as its one and only element. [0 points]

<!-- grade_id: cell-3af2459c61b6c472, points: 0, grade: false, solution: true -->
```python
def itself():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6dbaa09e08641787, points: 0, grade: true, solution: false -->
```python
list_ = itself()
assert list_ is list_[0]
```

---

**BONUS Q**: Create a function that returns the first positive integer for which Python does not create a singleton. What about the negative integer closest to zero? [0 points]

<!-- grade_id: cell-6f599029a544dfcd, points: 0, grade: false, solution: true -->
```python
def first_non_singleton():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8a115bd7fb268cc5, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f5077f7c620f08cb, points: 0, grade: false, solution: true -->
```python
def last_negative_non_singleton():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4ea5e3cac5553b11, points: 0, grade: true, solution: false -->
```python

```