**Q**: Create a generator function `fibonacci` that returns a generator object that can be used to produce the Fibonacci sequence. Your function should be able to generate the 10,000th Fibonacci number in under a few seconds. [2 points]

Please note that the Fibonacci sequence in this problem starts with 0.

```python
>>> fib = fibonacci()
>>> [next(fib) for _ in range(5)]
[0, 1, 1, 2, 3]
```
[2 points]

<!-- grade_id: cell-32574a3fb8ed5982, points: 0, grade: false, solution: true -->
```python
def fibonacci():
    a = 0
    b = 1
    while 1:
        yield a
        a,b = b,a+b
```

<!-- grade_id: cell-2970ff4efe5245f7, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
fib = fibonacci()
assert [next(fib) for _ in range(5)] == [0, 1, 1, 2, 3]
### END TESTS
```

<!-- grade_id: cell-447405b27b670689, points: 0.5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-c4da759e55299384, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
import inspect

assert inspect.isgeneratorfunction(fibonacci)
### END TESTS
```

---

**Q**: Create a Markov chain generator function named `markov`. The generator should take a starting point (an integer). With each pass of the resultant generator object to `next`, a random step from the last point returned (or the starting point if no point has yet been returned) should be performed. The result of this step should be returned after the step is taken. The steps can be -1, 0, or 1 unit from the last point returned (or the starting point if no point has yet been returned).

[2 points]

<!-- grade_id: cell-e8a9dfd93b0e0309, points: 0, grade: false, solution: true -->
```python
import random

def markov(start: int):
    m = start
    count = 0
    yield start
    count+=1
    while count >0:
        m += random.randint(-1,1)
        yield m
```

<!-- grade_id: cell-59f5d59626ff94cd, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
import inspect

assert inspect.isgeneratorfunction(markov)
### END TESTS
```

<!-- grade_id: cell-c8b67e7d41fc06c4, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
position = 0
gen = markov(position)
 
for _ in range(1000):
    new_position = next(gen)
    assert abs(new_position - position) <= 1, f'{new_position} and {position} are not within 1 unit of each other'
    position = new_position
### END TESTS
```

<!-- grade_id: cell-6cb1eeb3900e89eb, points: 0.5, grade: true, solution: false -->
```python

```